#### PROBLEMA 14
    Write a function quantum_addition(a, b) that performs quantum addition of two integers, a and b. The function should:
    - Prepare the input quantum state based on the binary representation of a, with qubit_0 as the least significant bit (LSB) and qubit_n as the most significant bit (MSB), where qubit_0 is the first qubit in the circuit, as conventionally represented.
    - Construct a quantum circuit that adds b to a, ensuring the resulting quantum state encodes the binary representation of a+ba + ba+b.
    - Combine all operations into a single gate and return the resulting gate as a quantum circuit.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-486f064b05395e04, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit import QuantumCircuit, transpile
import matplotlib.pyplot as plt

#inverse encoding
def prepare_binary_state(a):

    n = a.bit_length()
    qc = QuantumCircuit(n*2)
    for i in range(n/2-1, -1, -1):
        if a>=2**i:
            qc.x(i)
            a -= 2**i

    return Operator(qc)


def quantum_addition(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d322040f709992b4, points: 75, grade: true, solution: false -->
```python
exp1 = [[(-6.938893903907227e-18-3.4694469519536136e-17j),0j,(6.938893903907227e-17+0j),0j,(-6.938893903907227e-18-3.4694469519536136e-17j),(-0.9238795325112864-0.3826834323650895j),-1.3877787807814454e-17j,0j,],[(-1.3877787807814457e-17+5.74836792573273e-18j),(4.212649542996074e-17-5.4320625253083316e-17j),(-0.9238795325112864-0.3826834323650896j),(6.413370311985692e-17-6.938893903907227e-18j),(1.3877787807814457e-17+3.3503943541361643e-17j),(-5.600428323777519e-17-1.5068313785988955e-17j),8.326672684688674e-17j,(5.255235919215356e-18-6.938893903907227e-18j),],[(3.4694469519536136e-17-6.938893903907227e-18j),0j,(-4.1633363423443364e-17-4.1633363423443364e-17j),(-0.9238795325112864-0.3826834323650896j),(3.4694469519536136e-17-6.938893903907227e-18j),0j,(2.7755575615628907e-17-2.7755575615628907e-17j),1.3877787807814457e-16j,],[(-0.9238795325112863-0.3826834323650896j),(8.37598588534041e-17-2.656504963745441e-17j),1.9626155733547187e-17j,(-3.637812750422801e-17-3.4694469519536136e-17j),(-1.1102230246251565e-16+1.3877787807814457e-16j),(-1.437091981433183e-17+1.2687261829639954e-17j),1.9626155733547187e-17j,(2.250033969641355e-17-3.4694469519536136e-17j),],[(-6.938893903907227e-18-3.4694469519536136e-17j),(-0.9238795325112864-0.3826834323650895j),-1.3877787807814454e-17j,0j,(-6.938893903907227e-18-3.4694469519536136e-17j),0j,(6.938893903907227e-17+0j),0j,],[(1.3877787807814457e-17+3.3503943541361643e-17j),(-5.600428323777519e-17-1.5068313785988955e-17j),8.326672684688674e-17j,(5.255235919215356e-18-6.938893903907227e-18j),(-1.3877787807814457e-17+5.74836792573273e-18j),(4.212649542996074e-17-5.4320625253083316e-17j),(-0.9238795325112864-0.3826834323650896j),(6.413370311985692e-17-6.938893903907227e-18j),],[(3.4694469519536136e-17-6.938893903907227e-18j),0j,(2.7755575615628907e-17-2.7755575615628907e-17j),1.3877787807814457e-16j,(3.4694469519536136e-17-6.938893903907227e-18j),0j,(-4.1633363423443364e-17-4.1633363423443364e-17j),(-0.9238795325112864-0.3826834323650896j),],[(-1.1102230246251565e-16+1.3877787807814457e-16j),(-1.437091981433183e-17+1.2687261829639954e-17j),1.9626155733547187e-17j,(2.250033969641355e-17-3.4694469519536136e-17j),(-0.9238795325112863-0.3826834323650896j),(8.37598588534041e-17-2.656504963745441e-17j),1.9626155733547187e-17j,(-3.637812750422801e-17-3.4694469519536136e-17j)]]
out1 = quantum_addition(2,1)
assert np.allclose(out1, exp1), f"Test failed"
```

Output:
```
Cell In[11], line 2
      1 exp1 = [[(-6.938893903907227e-18-3.4694469519536136e-17j),0j,(6.938893903907227e-17+0j),0j,(-6.938893903907227e-18-3.4694469519536136e-17j),(-0.9238795325112864-0.3826834323650895j),-1.3877787807814454e-17j,0j,],[(-1.3877787807814457e-17+5.74836792573273e-18j),(4.212649542996074e-17-5.4320625253083316e-17j),(-0.9238795325112864-0.3826834323650896j),(6.413370311985692e-17-6.938893903907227e-18j),(1.3877787807814457e-17+3.3503943541361643e-17j),(-5.600428323777519e-17-1.5068313785988955e-17j),8.326672684688674e-17j,(5.255235919215356e-18-6.938893903907227e-18j),],[(3.4694469519536136e-17-6.938893903907227e-18j),0j,(-4.1633363423443364e-17-4.1633363423443364e-17j),(-0.9238795325112864-0.3826834323650896j),(3.4694469519536136e-17-6.938893903907227e-18j),0j,(2.7755575615628907e-17-2.7755575615628907e-17j),1.3877787807814457e-16j,],[(-0.9238795325112863-0.3826834323650896j),(8.37598588534041e-17-2.656504963745441e-17j),1.9626155733547187e-17j,(-3.637812750422801e-
```

```python
a = 4
b = 3
n = a.bit_length()
qc = QuantumCircuit(n*2 + 1)
for i in range(n//2 - 1, -1, -1):
    if a>=2**i:
        qc.x(i)
        a -= 2**i

Operator(qc)
print(qc)
```

Output:
```
┌───┐
q_0: ┤ X ├
     └───┘
q_1: ─────
          
q_2: ─────
          
q_3: ─────
          
q_4: ─────
          
q_5: ─────
          
q_6: ─────
```

```python
def apply_inverse_qft(circuit, num_qubits):
    
    
    # Reverse the order of the qubits
    for i in range(num_qubits // 2):
        circuit.swap(i, num_qubits - i - 1)
        # back iteration and negative angle for inverse
    for qubit in range(num_qubits - 1, -1, -1): 
        for control_qubit in range(num_qubits - 1, qubit, -1):
            angle = -np.pi / (2 ** (control_qubit - qubit))
            circuit.cp(angle, control_qubit, qubit)
        
        circuit.h(qubit)
```

```python
def apply_qft(circuit, num_qubits):
    for qubit in range(num_qubits):
   
        circuit.h(qubit)
        
        for control_qubit in range(qubit + 1, num_qubits):
            angle = np.pi / (2 ** (control_qubit - qubit))
            circuit.cp(angle, control_qubit, qubit)
    
    for i in range(num_qubits // 2):
        circuit.swap(i, num_qubits - i - 1)
```

```python
import numpy as np
from qiskit import QuantumCircuit, transpile
from qiskit.quantum_info import Statevector
from qiskit_aer import AerSimulator
import matplotlib.pyplot as plt
# Inputs
a = 4
b = 3
n = max(a.bit_length(), b.bit_length()) 
total_qubits = n * 2 + 1
  # |a>|0>|b> .

qc = QuantumCircuit(total_qubits)


for i in range(n - 1, -1, -1):
    if a >= 2**i:
        qc.x(i) 
        a -= 2**i

  
for i in range(n - 1, -1, -1):
    if b >= 2**i:
        qc.x(2 * n + 1 - 1 - i)  #
        b -= 2**i


# q23Q    
print(qc)

apply_qft(qc, n+1)
print(qc)

    

#for j in range(n):
#  for s in range(n):
#      # Calculate the angle for the CP gate
#      angle = 2 * np.pi / 2**(j + s - n)
      # Apply the CP gate
#      qc.cp(angle, j, n + 1 + s)


#apply_invere_qft(qc,n+1)

Operator(qc).to_matrix()

```

Output:
```
q_0: ─────
          
q_1: ─────
     ┌───┐
q_2: ┤ X ├
     └───┘
q_3: ─────
          
q_4: ─────
     ┌───┐
q_5: ┤ X ├
     ├───┤
q_6: ┤ X ├
     └───┘
     ┌───┐                                                                     »
q_0: ┤ H ├─■────────■─────────────■────────────────────────────────────────────»
     └───┘ │P(π/2)  │       ┌───┐ │                                            »
q_1: ──────■────────┼───────┤ H ├─┼────────■────────■───────────────────────X──»
     ┌───┐          │P(π/4) └───┘ │        │P(π/2)  │       ┌───┐           │  »
q_2: ┤ X ├──────────■─────────────┼────────■────────┼───────┤ H ├─■─────────X──»
     └───┘                        │P(π/8)           │P(π/4) └───┘ │P(π/2) ┌───┐»
q_3: ─────────────────────────────■─────────────────■─────────────■───────┤ H ├»
                                                                          └───┘»
q_4: ──────────────────────────────────────────────────────────────────────────»
     ┌───┐                          
```

```python

```

```python

```