# `create_W_state(shots, seed)`

### Task
Implement a function `create_W_state(shots, seed)` that generates a W state starting from the initial state ∣000⟩. The function should:

1. **Prepare the W State**: Use a quantum circuit to create the W state:
   $$(
   |W⟩ = \frac{1}{\sqrt{3}}(|100⟩ + |010⟩ + |001⟩)
   )$$

2. **Simulate the Circuit**: Use a quantum simulator to execute the circuit with the specified `shots` and `seed`.

3. **Return Results**: Output the measurement results (counts) from the simulation.

---

### Parameters
- **`shots`** (`int`): The number of measurements for the simulation.
- **`seed`** (`int`): The seed for the random number generator to ensure reproducibility.

---

### Functionality Requirements
1. **Initial State**: Start with the state ∣000⟩.
2. **Circuit Design**:
   - Apply gates to transform ∣000⟩ into the W state.
3. **Simulation**: Run the circuit using a quantum simulator.
4. **Return**: The measurement results (counts).

```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-04fd3fd43ebd8174, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit_aer import AerSimulator
from qiskit import QuantumCircuit, transpile
import matplotlib.pyplot as plt

def create_W_state(shots, seed):
     qc = QuantumCircuit(3)
 
     qc.ry(2 * np.arccos(1 / np.sqrt(3)), 0)
     qc.ch(0,1)
     qc.cx(1,2)
     qc.cx(0, 1)
     qc.x(0)
     qc.draw('mpl')
     qc.measure_all()
   
     plt.show()
     backend = AerSimulator()
    
     transpiled_qc = transpile(qc, backend)

     job = backend.run(transpiled_qc, shots=shots, seed_simulator=seed)

     result = job.result()
     counts = result.get_counts(transpiled_qc)
     print(counts)


     return counts
```

<!-- grade_id: cell-c86d340314477725, points: 10, grade: true, solution: false -->
```python
counts_expected = {'010': 3315, '100': 3429, '001': 3256}
cts = create_W_state(10000, 10)
assert cts == counts_expected, f"Test Failed"
```

Output:
```
{'001': 3256, '100': 3429, '010': 3315}
```

<!-- grade_id: cell-69ebaffbf2824e0f, points: 30, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-9cb74ace68b50d3f, points: 30, grade: true, solution: false -->
```python

```