### PROBLEMA 28
    Create a Python function xor_transform_gate(x, k, num_bits) that constructs a quantum circuit and returns a custom gate encapsulating all operations. The function should:
1. Prepare the quantum state $∣xk⟩$ by encoding the binary representations of x and k (padded with leading zeros to match num_bits).
2. Apply a quantum operation to transform the state $∣xk⟩$ into $∣x(xi⊕ki)⟩$, where $x_i⊕k_i$​ represents the XOR operation applied bitwise.
3. Combine all gates into a single Qiskit matrix representation that implements and returns the transformation.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-e5422bc6ebc127f9, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit import QuantumCircuit
from qiskit.quantum_info import Operator

def xor_transform_gate(x, k, num_bits):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fd819c5a44ba4125, points: 25, grade: true, solution: false -->
```python
exp1 = [[(1+0j),0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,],[0j,0j,0j,0j,0j,(1+0j),0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,],[0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,(1+0j),0j,0j,0j,0j,0j,],[0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,(1+0j),],[0j,0j,0j,0j,(1+0j),0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,],[0j,(1+0j),0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,],[0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,(1+0j),0j,],[0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,(1+0j),0j,0j,0j,0j,],[0j,0j,0j,0j,0j,0j,0j,0j,(1+0j),0j,0j,0j,0j,0j,0j,0j,],[0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,(1+0j),0j,0j,],[0j,0j,(1+0j),0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,],[0j,0j,0j,0j,0j,0j,0j,(1+0j),0j,0j,0j,0j,0j,0j,0j,0j,],[0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,(1+0j),0j,0j,0j,],[0j,0j,0j,0j,0j,0j,0j,0j,0j,(1+0j),0j,0j,0j,0j,0j,0j,],[0j,0j,0j,0j,0j,0j,(1+0j),0j,0j,0j,0j,0j,0j,0j,0j,0j,],[0j,0j,0j,(1+0j),0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j,0j]]
out1 = xor_transform_gate(3, 1, 2)
assert np.allclose(out1, exp1), f"Test failed"
```

<!-- grade_id: cell-0ca04beb5bc887f6, points: 25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-df40275dcb580065, points: 25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-ba3560fe5efbdf2f, points: 25, grade: true, solution: false -->
```python

```