#### PROBLEMA 10
Write a function quantum_fourier_transform(num_qubits, shots, seed, input_state) that creates and simulates a Quantum Fourier Transform (QFT) circuit. The function should:

- Initialize the qubits with the given input_state.
- Apply the QFT to the qubits.
- Perform the specified number of shots measurements.
- Use the provided seed for reproducibility.
- Include the final swap of the algorithm


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-8a80cacad296f511, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit import QuantumCircuit, transpile
from qiskit.quantum_info import Statevector
from qiskit_aer import AerSimulator
from qiskit.circuit.library import Initialize


def apply_qft(circuit, num_qubits):
    for qubit in range(num_qubits):
   
        circuit.h(qubit)
        
        for control_qubit in range(qubit + 1, num_qubits):
            angle = np.pi / (2 ** (control_qubit - qubit))
            circuit.cp(angle, control_qubit, qubit)
    
    for i in range(num_qubits // 2):
        circuit.swap(i, num_qubits - i - 1)


def quantum_fourier_transform(num_qubits, shots, seed, input_state):
    qc = QuantumCircuit(num_qubits)
    qc.initialize(Statevector(input_state), range(num_qubits))
    apply_qft(qc, num_qubits)
    qc.measure_all()
    
    backend = AerSimulator()
    
    transpiled_qc = transpile(qc, backend)

    job = backend.run(transpiled_qc, shots=shots, seed_simulator=seed)

    result = job.result()
    counts = result.get_counts(transpiled_qc)
    return counts
    



```

<!-- grade_id: cell-c39602bcffd90f2d, points: 50, grade: true, solution: false -->
```python
exp1 = {'1': 479, '0': 521}
state1 = [1, 0]
out1 = quantum_fourier_transform(1, 1000, 19, state1)
assert exp1 == out1, f"Test failed"
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[11], line 4
      2 state1 = [1, 0]
      3 out1 = quantum_fourier_transform(1, 1000, 19, state1)
----> 4 assert exp1 == out1, f"Test failed"

AssertionError: Test failed
```

<!-- grade_id: cell-075267e77f26a616, points: 50, grade: true, solution: false -->
```python
num_qubits = 3
input_state = [1, 0,0,0,0,0,0,0]
qc = QuantumCircuit(num_qubits)

qc.initialize(Statevector(input_state), range(num_qubits))
apply_qft(qc, num_qubits-1)
qc.measure_all()

backend = AerSimulator()

transpiled_qc = transpile(qc, backend)

job = backend.run(transpiled_qc, shots=1000, seed_simulator=19)

result = job.result()
counts = result.get_counts(transpiled_qc)
print(counts)
```

Output:
```
{'001': 255, '010': 227, '000': 266, '011': 252}
```

<!-- grade_id: cell-0ea0d5e6d080c95e, points: 50, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-1f1ffcdb0e186e04, points: 50, grade: true, solution: false -->
```python

```