#### PROBLEMA 17
    Write a function phase_flip_code(input_state, error) that implements a repetition code for detecting phase-flip errors. The function should:
        - Encode a single-qubit quantum state (input_state) into a three-qubit (repetition code) state using a stabilizer code.
        - Apply phase-flip errors (Z gates) based on the qubit indices provided in the error list.
        - Decode the three-qubit state and return the error syndrome after the decoding process.

    The function should use stabilizer code principles to detect phase-flip errors. HINT. The circuit can be constructed starting from the previous one (problem 16) and adapt the way the Z gate can be simplified.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-b2985341360b7b9d, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit import QuantumCircuit, transpile
from qiskit.quantum_info import Statevector
from qiskit_aer import AerSimulator

def phase_flip_code(input_state, error_qubits):

    qc = QuantumCircuit(3, 2)
    
    qc.initialize(Statevector(input_state), [0])
    
    qc.h(0)
    qc.cx(0, 1)
    qc.cx(0, 2)
    qc.h(0)
    qc.h(1)
    qc.h(2)

    for qubit in error_qubits:
        qc.z(qubit)

    qc.h(0)
    qc.h(1)
    qc.h(2)
    qc.cx(1, 2)
    qc.cx(0, 1)
    qc.measure([1, 2], [0, 1])
    

    qc.cx(0, 2)
    qc.cx(0, 1)
    qc.h(0)
    qc.h(1)
    qc.h(2)
    
    # Simulate the circuit
    backend = AerSimulator()
    transpiled_qc = transpile(qc, backend)
    result = backend.run(transpiled_qc, shots=1).result()
    counts = result.get_counts(transpiled_qc)
    return counts
```

<!-- grade_id: cell-5e53cca07d23857a, points: 40, grade: true, solution: false -->
```python
exp1 = '01'
input_state = [1/np.sqrt(2), 1/np.sqrt(2)]
out1 = phase_flip_code(input_state, [0])
assert exp1 == out1, f"Test failed"
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[7], line 4
      2 input_state = [1/np.sqrt(2), 1/np.sqrt(2)]
      3 out1 = phase_flip_code(input_state, [0])
----> 4 assert exp1 == out1, f"Test failed"

AssertionError: Test failed
```

```python
    error_qubits = [0]    
    qc = QuantumCircuit(3, 2)
    
    qc.initialize(Statevector(input_state), [0])
    
    qc.h(0)
    qc.cx(0, 1)
    qc.cx(0, 2)
    qc.h(0)
    qc.h(1)
    qc.h(2)

    for qubit in error_qubits:
        qc.z(qubit)

    qc.h(0)
    qc.h(1)
    qc.h(2)
    qc.cx(1, 2)
    qc.cx(0, 1)
    qc.measure([1, 2], [0, 1])
    

    qc.cx(0, 2)
    qc.cx(0, 1)
    qc.h(0)
    qc.h(1)
    qc.h(2)
    
    # Simulate the circuit
    backend = AerSimulator()
    transpiled_qc = transpile(qc, backend)
    result = backend.run(transpiled_qc, shots=1).result()
    counts = result.get_counts(transpiled_qc)
    print(counts)
```

Output:
```
{'01': 1}
```