#### PROBLEMA 15
    Write a function quantum_modular_addition(a, b, N) that performs modular addition of two integers a, b mod N. The function should:
        - Prepare the input quantum state based on the binary representation of aaa, with qubit_0 as the least significant bit (LSB) and qubit_n as the most significant bit (MSB), where qubit_0 is the first qubit in the circuit, as conventionally represented.
        - Construct a quantum circuit that computes (a+b)mod  N, ensuring the resulting quantum state encodes the binary representation of the modular sum.
        - Combine all operations into a single gate and return the resulting gate as a quantum circuit.
    Note: Ensure that a and b are less than N.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-4cd4844898c26a65, points: 0, grade: false, solution: true -->
```python
from qiskit import QuantumCircuit, transpile
import numpy as np
from qiskit_aer import AerSimulator
from qiskit.quantum_info import Operator


def quantum_modular_addition(a, b, N):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-0065f93c7e0bb946, points: 100, grade: true, solution: false -->
```python
exp1 = [[(4.1633363423443364e-17-8.73314367879276e-17j),(0.7071067811865466-0.7071067811865467j),(-1.5602298185534274e-16+6.77052810543804e-17j),(-1.9626155733547187e-17-1.9626155733547187e-17j),(-4.037530741417945e-18-9.8402570663965e-18j),(-5.5511151231257784e-17-5.5511151231257784e-17j),(-5.236017322854745e-18-5.193957329673427e-17j),(-2.775557561562889e-17-1.9595346869773345e-16j),(3.9252311467094367e-17+1.556144579250632e-17j),(-5.551115123125783e-17+0j),(-3.75686534824025e-17+9.813077866773593e-18j),(9.813077866773593e-17+3.925231146709437e-17j),(-4.3810153414652586e-17-1.3365437204936483e-17j),(-3.5526689213701746e-17-3.5526689213701746e-17j),4.622231866529366e-33j,1.5407439555097887e-32j,],[(0.7071067811865468-0.7071067811865468j),(-2.943923360032078e-17+3.182028555666977e-17j),(-8.274848601723403e-17-2.0502702847208728e-17j),(-7.107259702376414e-17-7.443991299314789e-17j),(-1.1102230246251556e-16-6.162975822039155e-33j),(1.6836579846918662e-18-1.5561445792506332e-17j),(2.1847273144252364e-17-1.1216850102119418e-16j),(7.512371736499353e-17-1.0047902242791279e-16j),5.551115123125783e-17j,(-1.1496735851465466e-17-3.367315969383745e-18j),(1.1744295553677017e-16+5.582513009577077e-17j),(-2.7755575615628904e-17+8.129419882081724e-18j),(-1.4939806681594836e-16+8.611580198661776e-17j),(-2.8685549663172185e-18-3.468087991972465e-17j),(-3.0814879110195774e-33-4.4296388720906425e-33j),(-3.0814879110195774e-33-7.703719777548943e-33j),],[(1.2083538032928923e-16-2.77555756156289e-17j),(7.850462293418875e-17+1.9626155733547187e-17j),(-5.719480921594969e-17-3.1122891585012646e-17j),(0.7071067811865466+0.7071067811865466j),(-5.43206252530833e-17-2.822948902578567e-17j),(3.559071334380683e-33+2.609113622729763e-17j),(-4.037530741417955e-18-1.7915318549232362e-17j),(1.1102230246251558e-16-1.1102230246251556e-16j),(-5.719480921594969e-17+2.381051956348994e-18j),(-7.850462293418875e-17+1.9626155733547183e-17j),(3.182028555666977e-17+5.144644129021696e-17j),(-1.1102230246251565e-16-5.551115123125783e-17j),(-1.232595164407831e-32+1.232595164407831e-32j),(-6.162975822039155e-33+2.311115933264683e-33j),(5.072982872219922e-17+6.294317714168233e-17j),(-3.2517679528326883e-17+1.895269253967043e-16j),],[(1.0713674566347921e-16+9.864901949738865e-17j),(7.682096494949687e-17+2.369086567458805e-17j),(0.7071067811865468+0.7071067811865468j),(-2.3690865674588056e-17+7.107259702376414e-17j),(-2.997669302633407e-17-5.908302471376568e-18j),(3.3616869728346092e-18+1.9639745333358644e-17j),(1.1102230246251552e-16-1.6653345369377338e-16j),(-4.0647099410408296e-18+1.6836579846918782e-18j),(-8.274848601723404e-17-8.274848601723404e-17j),(-1.1438961843189937e-16-6.700788708272327e-17j),(-1.1102230246251565e-16+0j),(5.887846720064155e-17+1.4499971309929962e-32j),(-6.162975822039155e-33+1.232595164407831e-32j),(-3.0814879110195774e-33+6.162975822039155e-33j),(-2.5348147537455674e-16+1.7967606571990575e-17j),(-1.1510325451276935e-17+1.0516100834979847e-17j),],[(4.1633363423443364e-17+3.182028555666978e-17j),0j,(4.50006793928271e-17+1.2658374825502195e-16j),(-9.813077866773593e-17+1.373830901348303e-16j),(-1.1022822441308676e-17+1.5068313785988925e-17j),(-4.774003763318494e-17+4.774003763318493e-17j),(3.158525108795067e-32-3.0814879110195774e-32j),(3.235562306570556e-32-1.232595164407831e-32j),(-9.476346269835218e-17+1.2194129823122587e-17j),(0.7071067811865466+0.7071067811865467j),(-7.345364898011313e-17-3.75686534824025e-17j),(1.9626155733547187e-17+7.850462293418875e-17j),(9.840257066396502e-18-4.037530741417957e-18j),(5.5511151231257765e-17-5.5511151231257765e-17j),(-4.06470994104087e-18+8.148638478442335e-18j),(1.1268674185084684e-16-2.7755575615628898e-17j),],[(-5.551115123125783e-17+0j),(2.943923360032078e-17+5.144644129021696e-17j),(-3.5008448384049087e-17+1.3825963724849183e-16j),(-1.2194129823122584e-17+1.021954886087768e-16j),(-9.087790589931572e-17+1.237328296512701e-17j),(-3.113648118482412e-17+2.4393888642794308e-17j),(-4.622231866529366e-33+1.232595164407831e-32j),(6.162975822039155e-33+2.465190328815662e-32j),(0.7071067811865468+0.7071067811865468j),(6.700788708272329e-17+1.1438961843189937e-16j),(3.139788645129449e-19+6.193180430551235e-17j),(-2.775557561562891e-17+4.738173134917609e-17j),(-1.232595164407831e-32-1.1102230246251559e-16j),(1.5561445792506326e-17+1.6836579846918662e-18j),(1.8595564301016526e-17-2.336724536853082e-18j),(-2.778838381180098e-17-2.5388113259091364e-17j),],[(-3.7568653482402494e-17-2.7755575615628914e-17j),(3.925231146709437e-17+1.9626155733547187e-17j),(1.6836579846918662e-18+3.3673159693837356e-18j),(5.551115123125783e-17+0j),(-1.232595164407831e-32-9.244463733058732e-33j),(5.156692657724563e-33+0j),(-3.231341756318713e-17-4.4488328789949095e-17j),(1.0626019854981759e-16-5.0749047318559783e-17j),(5.719480921594969e-17+2.5374523659279913e-17j),(-3.925231146709437e-17-1.9626155733547183e-17j),(2.3690865674588047e-17-2.3690865674588047e-17j),(0.7071067811865467-0.7071067811865466j),(-1.7037468307331387e-16+1.131990924537248e-16j),(-5.717559061958903e-17+1.3877787807814438e-17j),(-1.7915318549232353e-17+4.0375307414179424e-18j),(-1.110223024625155e-16-1.110223024625155e-16j),],[(-1.0713674566347921e-16-2.0144396563199907e-17j),(-2.130981371823906e-17+4.064709941040862e-18j),(5.551115123125783e-17+0j),(-4.064709941040857e-18+1.2194129823122578e-17j),(3.0814879110195774e-33+2.6963019221421302e-33j),(-9.822242716374903e-33-6.162975822039155e-33j),(1.0103007145373754e-16+1.010300714537376e-16j),(5.62899654913711e-21-5.734778325921217e-18j),(-3.5008448384049074e-17-3.5008448384049074e-17j),(3.3673159693837325e-18+1.149673585146546e-17j),(0.7071067811865467-0.7071067811865468j),(-5.887846720064155e-17-1.0142089313877339e-32j),(-5.329003382055265e-17+6.10995846113467e-17j),(-4.6177615771190116e-17-3.062413058194613e-17j),(-1.665334536937733e-16-1.1102230246251558e-16j),(1.6836579846918686e-18+4.064709941040868e-18j),],[(9.840257066396502e-18-4.037530741417957e-18j),(5.5511151231257765e-17-5.5511151231257765e-17j),(-4.06470994104087e-18+8.148638478442335e-18j),(1.1268674185084684e-16-2.7755575615628898e-17j),(-9.476346269835218e-17+1.2194129823122587e-17j),(0.7071067811865466+0.7071067811865467j),(-7.345364898011313e-17-3.75686534824025e-17j),(1.9626155733547187e-17+7.850462293418875e-17j),(-1.1022822441308676e-17+1.5068313785988925e-17j),(-4.774003763318494e-17+4.774003763318493e-17j),(3.158525108795067e-32-3.0814879110195774e-32j),(3.235562306570556e-32-1.232595164407831e-32j),(4.1633363423443364e-17+3.182028555666978e-17j),0j,(4.50006793928271e-17+1.2658374825502195e-16j),(-9.813077866773593e-17+1.373830901348303e-16j),],[(-1.232595164407831e-32-1.1102230246251559e-16j),(1.5561445792506326e-17+1.6836579846918662e-18j),(1.8595564301016526e-17-2.336724536853082e-18j),(-2.778838381180098e-17-2.5388113259091364e-17j),(0.7071067811865468+0.7071067811865468j),(6.700788708272329e-17+1.1438961843189937e-16j),(3.139788645129449e-19+6.193180430551235e-17j),(-2.775557561562891e-17+4.738173134917609e-17j),(-9.087790589931572e-17+1.237328296512701e-17j),(-3.113648118482412e-17+2.4393888642794308e-17j),(-4.622231866529366e-33+1.232595164407831e-32j),(6.162975822039155e-33+2.465190328815662e-32j),(-5.551115123125783e-17+0j),(2.943923360032078e-17+5.144644129021696e-17j),(-3.5008448384049087e-17+1.3825963724849183e-16j),(-1.2194129823122584e-17+1.021954886087768e-16j),],[(-1.7037468307331387e-16+1.131990924537248e-16j),(-5.717559061958903e-17+1.3877787807814438e-17j),(-1.7915318549232353e-17+4.0375307414179424e-18j),(-1.110223024625155e-16-1.110223024625155e-16j),(5.719480921594969e-17+2.5374523659279913e-17j),(-3.925231146709437e-17-1.9626155733547183e-17j),(2.3690865674588047e-17-2.3690865674588047e-17j),(0.7071067811865467-0.7071067811865466j),(-1.232595164407831e-32-9.244463733058732e-33j),(5.156692657724563e-33+0j),(-3.231341756318713e-17-4.4488328789949095e-17j),(1.0626019854981759e-16-5.0749047318559783e-17j),(-3.7568653482402494e-17-2.7755575615628914e-17j),(3.925231146709437e-17+1.9626155733547187e-17j),(1.6836579846918662e-18+3.3673159693837356e-18j),(5.551115123125783e-17+0j),],[(-5.329003382055265e-17+6.10995846113467e-17j),(-4.6177615771190116e-17-3.062413058194613e-17j),(-1.665334536937733e-16-1.1102230246251558e-16j),(1.6836579846918686e-18+4.064709941040868e-18j),(-3.5008448384049074e-17-3.5008448384049074e-17j),(3.3673159693837325e-18+1.149673585146546e-17j),(0.7071067811865467-0.7071067811865468j),(-5.887846720064155e-17-1.0142089313877339e-32j),(3.0814879110195774e-33+2.6963019221421302e-33j),(-9.822242716374903e-33-6.162975822039155e-33j),(1.0103007145373754e-16+1.010300714537376e-16j),(5.62899654913711e-21-5.734778325921217e-18j),(-1.0713674566347921e-16-2.0144396563199907e-17j),(-2.130981371823906e-17+4.064709941040862e-18j),(5.551115123125783e-17+0j),(-4.064709941040857e-18+1.2194129823122578e-17j),],[(-4.3810153414652586e-17-1.3365437204936483e-17j),(-3.5526689213701746e-17-3.5526689213701746e-17j),4.622231866529366e-33j,1.5407439555097887e-32j,(3.9252311467094367e-17+1.556144579250632e-17j),(-5.551115123125783e-17+0j),(-3.75686534824025e-17+9.813077866773593e-18j),(9.813077866773593e-17+3.925231146709437e-17j),(-4.037530741417945e-18-9.8402570663965e-18j),(-5.5511151231257784e-17-5.5511151231257784e-17j),(-5.236017322854745e-18-5.193957329673427e-17j),(-2.775557561562889e-17-1.9595346869773345e-16j),(4.1633363423443364e-17-8.73314367879276e-17j),(0.7071067811865466-0.7071067811865467j),(-1.5602298185534274e-16+6.77052810543804e-17j),(-1.9626155733547187e-17-1.9626155733547187e-17j),],[(-1.4939806681594836e-16+8.611580198661776e-17j),(-2.8685549663172185e-18-3.468087991972465e-17j),(-3.0814879110195774e-33-4.4296388720906425e-33j),(-3.0814879110195774e-33-7.703719777548943e-33j),5.551115123125783e-17j,(-1.1496735851465466e-17-3.367315969383745e-18j),(1.1744295553677017e-16+5.582513009577077e-17j),(-2.7755575615628904e-17+8.129419882081724e-18j),(-1.1102230246251556e-16-6.162975822039155e-33j),(1.6836579846918662e-18-1.5561445792506332e-17j),(2.1847273144252364e-17-1.1216850102119418e-16j),(7.512371736499353e-17-1.0047902242791279e-16j),(0.7071067811865468-0.7071067811865468j),(-2.943923360032078e-17+3.182028555666977e-17j),(-8.274848601723403e-17-2.0502702847208728e-17j),(-7.107259702376414e-17-7.443991299314789e-17j),],[(-1.232595164407831e-32+1.232595164407831e-32j),(-6.162975822039155e-33+2.311115933264683e-33j),(5.072982872219922e-17+6.294317714168233e-17j),(-3.2517679528326883e-17+1.895269253967043e-16j),(-5.719480921594969e-17+2.381051956348994e-18j),(-7.850462293418875e-17+1.9626155733547183e-17j),(3.182028555666977e-17+5.144644129021696e-17j),(-1.1102230246251565e-16-5.551115123125783e-17j),(-5.43206252530833e-17-2.822948902578567e-17j),(3.559071334380683e-33+2.609113622729763e-17j),(-4.037530741417955e-18-1.7915318549232362e-17j),(1.1102230246251558e-16-1.1102230246251556e-16j),(1.2083538032928923e-16-2.77555756156289e-17j),(7.850462293418875e-17+1.9626155733547187e-17j),(-5.719480921594969e-17-3.1122891585012646e-17j),(0.7071067811865466+0.7071067811865466j),],[(-6.162975822039155e-33+1.232595164407831e-32j),(-3.0814879110195774e-33+6.162975822039155e-33j),(-2.5348147537455674e-16+1.7967606571990575e-17j),(-1.1510325451276935e-17+1.0516100834979847e-17j),(-8.274848601723404e-17-8.274848601723404e-17j),(-1.1438961843189937e-16-6.700788708272327e-17j),(-1.1102230246251565e-16+0j),(5.887846720064155e-17+1.4499971309929962e-32j),(-2.997669302633407e-17-5.908302471376568e-18j),(3.3616869728346092e-18+1.9639745333358644e-17j),(1.1102230246251552e-16-1.6653345369377338e-16j),(-4.0647099410408296e-18+1.6836579846918782e-18j),(1.0713674566347921e-16+9.864901949738865e-17j),(7.682096494949687e-17+2.369086567458805e-17j),(0.7071067811865468+0.7071067811865468j),(-2.3690865674588056e-17+7.107259702376414e-17j)]]
out1 = quantum_modular_addition(3,2,4)
assert np.allclose(out1, exp1), f"Test failed"
```

<!-- grade_id: cell-27857a1e4abf404e, points: 100, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-f1b080ed3949abfa, points: 100, grade: true, solution: false -->
```python

```