### PROBLEMA 30
    Create a function verify_orthonormality(basis_vectors) that accepts a list of basis vectors and determines whether they form an orthonormal set. The function should return True if the vectors are orthonormal; otherwise, return False

```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-cc9a80bd967d96e0, points: 0, grade: false, solution: true -->
```python
import numpy as np

def verify_orthonormality(basis_vectors):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c756b41fec1eda20, points: 20, grade: true, solution: false -->
```python
exp1 = True
basis_vectors1 = [np.array([1, 0, 0]), np.array([0, 1, 0]), np.array([0, 0, 1])]
out1 = verify_orthonormality(basis_vectors1)
assert exp1 == out1, f"Test failed"
```

<!-- grade_id: cell-42c21665b31d3842, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-ba69e5814afbfd85, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-ee941db2e6e11d65, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-aab021b81de045aa, points: 20, grade: true, solution: false -->
```python

```