### PROBLEMA 25
    Write a function custom_measurement_circuit that takes the following arguments: input_state (a state in the computational basis), observable (a Hermitian operator to measure), seed (to initialize a quantum simulator), and shots (the number of measurement runs). The function should construct a quantum circuit that measures the input_state in the eigenbasis of the observable. It should then return the expected value of the measurement based on the simulator's initialization and the specified number of shots.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-20afaa1c6d11cb65, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit import QuantumCircuit, transpile
from qiskit.quantum_info import Statevector
from qiskit_aer import AerSimulator
from qiskit.circuit.library import UnitaryGate


def custom_measurement_circuit(input_state, observable, seed, shots):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-ab0b02526e637719, points: 10, grade: true, solution: false -->
```python
exp1 = -0.30319999999999997
Z_obs = np.array([[1, 0], [0, -1]])
state1 = np.array([0.6, +0.8j])
out1 = custom_measurement_circuit(state1, Z_obs, 10, 5000)
assert np.allclose(out1, exp1), f"Test failed"
```

<!-- grade_id: cell-ebb86b4512fa31e3, points: 18, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-0395004221150776, points: 18, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-11c8e4fd40179a25, points: 18, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-b10623d7e5db4de8, points: 18, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-1901b5010b5a8fb9, points: 18, grade: true, solution: false -->
```python

```