#### PROBLEMA 12
    Write a function qft_then_inverse_qft(num_qubits, input_state) that constructs a quantum circuit for a given number of qubits (num_qubits) and an input_state represented as a list or array of complex coefficients. The circuit should apply the Quantum Fourier Transform (QFT) followed by its inverse to the input_state. The function should return the final state vector after these operations.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-93832c5ab582d6f1, points: 0, grade: false, solution: true -->
```python
from qiskit import QuantumCircuit
from qiskit.quantum_info import Statevector
import numpy as np

def apply_inverse_qft(circuit, num_qubits):
    
    
    # Reverse the order of the qubits
    for i in range(num_qubits // 2):
        circuit.swap(i, num_qubits - i - 1)
        # back iteration and negative angle for inverse
    for qubit in range(num_qubits - 1, -1, -1): 
        for control_qubit in range(num_qubits - 1, qubit, -1):
            angle = -np.pi / (2 ** (control_qubit - qubit))
            circuit.cp(angle, control_qubit, qubit)
        
        circuit.h(qubit)

def apply_qft(circuit, num_qubits):
    for qubit in range(num_qubits):
   
        circuit.h(qubit)
        
        for control_qubit in range(qubit + 1, num_qubits):
            angle = np.pi / (2 ** (control_qubit - qubit))
            circuit.cp(angle, control_qubit, qubit)
    
    for i in range(num_qubits // 2):
        circuit.swap(i, num_qubits - i - 1)

def qft_then_inverse_qft(num_qubits, input_state):
    qc = QuantumCircuit(num_qubits)
    qc.initialize(Statevector(input_state), range(num_qubits))
    apply_qft(qc, num_qubits)
    apply_inverse_qft(qc, num_qubits)
    final_state = Statevector.from_instruction(qc)
    return final_state


```

<!-- grade_id: cell-aebc550b7700cd95, points: 50, grade: true, solution: false -->
```python
exp1 = [0.5+0.j, 0.5+0.j, 0.5+0.j, 0.5+0.j]
input_state1 = [1/2, 1/2, 1/2, 1/2]
out1 = qft_then_inverse_qft(2, input_state1).data
assert np.allclose(out1, exp1), f"Test failed"
```

<!-- grade_id: cell-8caaf115a91f0e68, points: 50, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-bae72a23805f29c3, points: 50, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-324df13cdaef9472, points: 50, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-8bfda5843ae31e62, points: 50, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-05540d195c01215d, points: 50, grade: true, solution: false -->
```python

```