### PROBLEMA 23
    Write a function create_cu3_gate(theta, phi, lambda, num_qubits, control_qubit, target_qubit) that generates the matrix representation of a custom controlled-U3 (CU3) gate. Return the matrix of the gate. Use the normal way of representing qubits (not the one Qiskit does).

```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-153e521616421585, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit import QuantumCircuit
from qiskit.quantum_info import Operator
from qiskit.circuit.library import UnitaryGate

def create_cu3_gate(theta, phi, lamda, num_qubits, control_qubit, target_qubit):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1e56149605ad40ae, points: 25, grade: true, solution: false -->
```python
exp1 = [[(0.9999999999999999-6.106226635438361e-16j),0j,(5.551115123125783e-17+9.614813431917822e-17j),0j,0j,0j,0j,0j,],[0j,(0.923879532511287+5.551115123125783e-17j),0j,(-0.331413574035592-0.1913417161825446j),0j,0j,0j,0j,],[(2.7755575615628914e-17+4.807406715958911e-17j),0j,(0.9999999999999997-8.881784197001252e-16j),0j,0j,0j,0j,0j,],[0j,(0.1913417161825451+0.3314135740355918j),0j,(6.661338147750939e-16+0.9238795325112867j),0j,0j,0j,0j,],[0j,0j,0j,0j,(0.9999999999999999-6.106226635438361e-16j),0j,(5.551115123125783e-17+9.614813431917822e-17j),0j,],[0j,0j,0j,0j,0j,(0.923879532511287+5.551115123125783e-17j),0j,(-0.331413574035592-0.1913417161825446j),],[0j,0j,0j,0j,(2.7755575615628914e-17+4.807406715958911e-17j),0j,(0.9999999999999997-8.881784197001252e-16j),0j,],[0j,0j,0j,0j,0j,(0.1913417161825451+0.3314135740355918j),0j,(6.661338147750939e-16+0.9238795325112867j)]]
theta, phi, lamda = np.pi/4, np.pi/3, np.pi/6
out1 = create_cu3_gate(theta, phi, lamda, 3, 0, 1)
assert np.allclose(exp1, out1), f"Test failed"
```

<!-- grade_id: cell-4c104786fffc63ee, points: 25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-369a7e0868939b14, points: 25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-267385755d72d536, points: 25, grade: true, solution: false -->
```python

```