Write a function measure_state(state_vector, shots, seed) that simulates the measurement of a quantum state. The input state_vector is a list or array of complex coefficients representing the state, shots is the number of measurements to perform, and seed ensures reproducibility. The function should simulate the specified number of measurements and return a list of individual measurement outcomes (memory) from the simulation.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-25f99054d158d002, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit_aer import AerSimulator
from qiskit import QuantumCircuit, transpile
from qiskit.quantum_info import Statevector


def measure_state(state_vector, shots, seed):

    # init quantum circuit of a given state vector
    qc = QuantumCircuit(1)
    qc.initialize(Statevector(state_vector), [0])
    qc.measure_all()

    backend = AerSimulator()
    
    transpiled_qc = transpile(qc, backend)

    job = backend.run(transpiled_qc, shots=shots, memory = True, seed_simulator=seed)

    result = job.result()
    memory = result.get_memory(transpiled_qc)
    print(memory)


    return memory
    
```

```python
state = (1/np.sqrt(2))*np.array([1, 1])
measure_state(state, 100, 199)
```

Output:
```
['0', '1', '1', '1', '0', '1', '1', '1', '0', '0', '1', '0', '1', '0', '1', '1', '1', '1', '1', '0', '1', '0', '0', '0', '1', '0', '0', '0', '0', '1', '1', '0', '0', '1', '1', '1', '1', '1', '1', '1', '1', '0', '0', '0', '0', '1', '1', '0', '0', '0', '0', '0', '0', '1', '1', '0', '1', '0', '0', '0', '0', '0', '1', '0', '1', '1', '0', '1', '0', '0', '1', '1', '1', '1', '1', '1', '0', '1', '1', '0', '0', '1', '1', '0', '1', '1', '1', '1', '0', '0', '1', '1', '0', '1', '1', '0', '0', '0', '0', '1']
['0',
 '1',
 '1',
 '1',
 '0',
 '1',
 '1',
 '1',
 '0',
 '0',
 '1',
 '0',
 '1',
 '0',
 '1',
 '1',
 '1',
 '1',
 '1',
 '0',
 '1',
 '0',
 '0',
 '0',
 '1',
 '0',
 '0',
 '0',
 '0',
 '1',
 '1',
 '0',
 '0',
 '1',
 '1',
 '1',
 '1',
 '1',
 '1',
 '1',
 '1',
 '0',
 '0',
 '0',
 '0',
 '1',
 '1',
 '0',
 '0',
 '0',
 '0',
 '0',
 '0',
 '1',
 '1',
 '0',
 '1',
 '0',
 '0',
 '0',
 '0',
 '0',
 '1',
 '0',
 '1',
 '1',
 '0',
 '1',
 '0',
 '0',
 '1',
 '1',
 '1',
 '1',
 '1',
 '1',
 '0',
 '1',
 '1',
 '0',
 '0',
 '1',
 '1',
 
```

<!-- grade_id: cell-c376bae88e144429, points: 10, grade: true, solution: false -->
```python
exp = ['0', '1', '1', '1', '0', '1', '1', '1', '0', '0', '1', '0', '1', '0', '1', '1', '1', '1', '1', '0', '1', '0', '0', '0', '1', '0', '0', '0', '0', '1', '1', '0', '0', '1', '1', '1', '1', '1', '1', '1', '1', '0', '0', '0', '0', '1', '1', '0', '0', '0', '0', '0', '0', '1', '1', '0', '1', '0', '0', '0', '0', '0', '1', '0', '1', '1', '0', '1', '0', '0', '1', '1', '1', '1', '1', '1', '0', '1', '1', '0', '0', '1', '1', '0', '1', '1', '1', '1', '0', '0', '1', '1', '0', '1', '1', '0', '0', '0', '0', '1']
state = (1/np.sqrt(2))*np.array([1, 1])
out = measure_state(state, 100, 199)
assert out == exp, f"Test failed"
```

Output:
```
['0', '1', '1', '1', '0', '1', '1', '1', '0', '0', '1', '0', '1', '0', '1', '1', '1', '1', '1', '0', '1', '0', '0', '0', '1', '0', '0', '0', '0', '1', '1', '0', '0', '1', '1', '1', '1', '1', '1', '1', '1', '0', '0', '0', '0', '1', '1', '0', '0', '0', '0', '0', '0', '1', '1', '0', '1', '0', '0', '0', '0', '0', '1', '0', '1', '1', '0', '1', '0', '0', '1', '1', '1', '1', '1', '1', '0', '1', '1', '0', '0', '1', '1', '0', '1', '1', '1', '1', '0', '0', '1', '1', '0', '1', '1', '0', '0', '0', '0', '1']
```

<!-- grade_id: cell-c7a015b181fd3cf0, points: 30, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-28e2932abac017e4, points: 30, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-6d4debeca732482c, points: 30, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-2fe4aae53081c160, points: 30, grade: true, solution: false -->
```python

```