# Quantum Machine Learning Challenge

**Challenge:** 

> Implement a function ```qnn(samples, labels, seed)``` that constructs a quantum circuit to train a simple quantum neural network (QNN) on the provided training data and returns the trained model.

**Details:**

1. Initialize the Quantum Circuit: Create a quantum circuit with enough qubits to encode the features of the data points.
2. Define Ansatz: Choose a simple parameterized quantum circuit as the ansatz for the QNN.
3. Encode Training Data: Use parameterized gates to encode the training data into the quantum circuit.
4. Train QNN: Use a classical optimizer to train the QNN by minimizing the loss function.
5. Output of QNN Function: Return the trained model.
6. Evaluate Model: Test the trained QNN on a validation dataset.


```python
!pip install scikit-learn
!pip install qiskit_machine_learning
!pip install matplotlib
```

```python
from sklearn.datasets import make_classification
from sklearn.model_selection import train_test_split 
from qiskit.circuit.library import *
from qiskit_machine_learning.optimizers import *
from qiskit.primitives import *
from qiskit_machine_learning.algorithms.classifiers import *
```

<!-- grade_id: cell-c0091def48cb0a50, points: 0, grade: false, solution: true -->
```python
def qnn(samples, labels, number_of_features):
    # Define the feature_map
    # Write your solution here | feature_map = ... 
    # YOUR CODE HERE
    raise NotImplementedError()

    # Create the ansatz
    # Write your solution here | ansatz = ... 
    # YOUR CODE HERE
    raise NotImplementedError()

    # Define the optimizer
    # Write your solution here | optimizer = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Create a sampler that is backed by a quantum computer
    # Write your solution here | sampler = ... 
    # YOUR CODE HERE
    raise NotImplementedError()

    # Create the quantum neural network (Veriational Quantum Classifier)
    # Write your solution here | vqc = ... 
    # YOUR CODE HERE
    raise NotImplementedError()

    # Fit the VQC model
    # YOUR CODE HERE
    raise NotImplementedError()

    return vqc
```

<!-- grade_id: cell-190cebe6ca50a149, points: 35, grade: true, solution: false -->
```python
"""Check that the function returns the correct output for several inputs"""
# Creating a dataset
number_of_samples = 100
number_of_features = 6
seed = 1234

X, y = make_classification(
    n_samples=number_of_samples, 
    n_features=number_of_features, n_informative=number_of_features,
    n_redundant=0, n_clusters_per_class=1, 
    class_sep=1.0, random_state=seed)
    
# Split the dataset in train and test using a test_size of 20% and use the seed as random_state
X_train, X_test, y_train, y_test = train_test_split(
    X, y , random_state=seed, test_size=0.2, shuffle=True)

QNN = qnn(X_train, y_train, number_of_features)
accuracy = QNN.score(X_test, y_test)

# Test the accuracy
assert accuracy > 0.5
```

<!-- grade_id: cell-b42cc2157e880f43, points: 35, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-5ec8ae68a368b1b8, points: 10, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-fd4165af7104503c, points: 20, grade: true, solution: false -->
```python

```