Write a function prepare_binary_state(a) that takes an integer a as input and creates a quantum circuit to prepare a quantum state representing the binary encoding of a. The binary representation should map qubit_0 to the most significant bit (MSB) and qubit_n to the least significant bit (LSB), with qubit_0 being the first qubit in the circuit as per standard notation. Combine all operations into a single gate and return the resulting gate as a quantum circuit.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-a757ae784930caa1, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit import QuantumCircuit
from qiskit.quantum_info import Operator

def prepare_binary_state(a):

    n = a.bit_length()
    qc = QuantumCircuit(n)  # if 1 put x gate
    for i in range(n):
        if a>=2**(n-i-1):
            qc.x(i)
            a -= 2**(n-i-1)

    return Operator(qc)

```

<!-- grade_id: cell-2d3c5b5985c7ad67, points: 25, grade: true, solution: false -->
```python
exp1 = [[0j,0j,0j,(1+0j),0j,0j,0j,0j,],[0j,0j,(1+0j),0j,0j,0j,0j,0j,],[0j,(1+0j),0j,0j,0j,0j,0j,0j,],[(1+0j),0j,0j,0j,0j,0j,0j,0j,],[0j,0j,0j,0j,0j,0j,0j,(1+0j),],[0j,0j,0j,0j,0j,0j,(1+0j),0j,],[0j,0j,0j,0j,0j,(1+0j),0j,0j,],[0j,0j,0j,0j,(1+0j),0j,0j,0j]]
out1 = prepare_binary_state(6)
assert np.allclose(out1, exp1), f"Test failed"
```

<!-- grade_id: cell-9c511dd89c146aaa, points: 25, grade: true, solution: false -->
```python

```