Write a function kronecker_product(state1, state2) that computes the Kronecker product of two quantum states, each represented as a list or array of complex coefficients. The function should return the resulting combined quantum state as an array of complex coefficients.


```python
!pip install numpy
```

<!-- grade_id: cell-e02eb809d484b0a4, points: 0, grade: false, solution: true -->
```python
import numpy as np

def kronecker_product(state1, state2):
   return( np.kron(state1, state2) )


    
```

<!-- grade_id: cell-3830ed149a93f445, points: 10, grade: true, solution: false -->
```python
exp = [0, 1, 0, 0]
state1_1 = np.array([1, 0])
state1_2 = np.array([0, 1])
out = kronecker_product(state1_1, state1_2)
assert np.allclose(exp, out), f"Test failed"
```

<!-- grade_id: cell-eeaae9e4d7dbe6db, points: 30, grade: true, solution: false -->
```python
state1_1 = np.array([1, 0])
state1_2 = np.array([0, 1])
print( np.kron(state1_1, state1_2) )
```

Output:
```
[0 1 0 0]
```

<!-- grade_id: cell-906ae59a68f045ec, points: 30, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-93c428f401f199c9, points: 30, grade: true, solution: false -->
```python

```