# Implementing a polynomial function with Classiq

**Challenge:** Implement a function ```polynomial(num_qbits_x, a)``` that creates a circuit to solve the equation $x^2 + a$. The input state x will be an equal superposition of ```num_qbits_x``` qubits and the function should return the results of the execution job.


```python
!pip install classiq
```

```python
from classiq import *
```

<!-- grade_id: cell-8d7d2a3743ead06b, points: 0, grade: false, solution: true -->
```python
def polynomial(num_qbts_x, a):
    # Define the circuit
    @qfunc
    def main(x: Output[QNum], y: Output[QNum]):
        # YOUR CODE HERE
        raise NotImplementedError()

    # Create the quantum program and execute the job
    # YOUR CODE HERE
    raise NotImplementedError()

    # Get the results
    results = job.result()[0].value.parsed_counts

    # Return the results
    return results
```

<!-- grade_id: cell-3a358e6996749bb5, points: 100, grade: true, solution: false -->
```python
"""Check that the function returns the correct output for several inputs"""
assert  "{'x': 0, 'y': 1}" in f'{polynomial(2,1)}'
assert  "{'x': 3, 'y': 10}" in f'{polynomial(2,1)}'
assert  "{'x': 7, 'y': 50}" in f'{polynomial(3,1)}'

```