# `create_EPRPair(input_state, shots, seed)`

### Task
Implement a function `create_EPRPair(input_state, shots, seed)` that performs the following steps:
1. **Initialize a Single-Qubit State**: Start with the given quantum state `input_state`.
2. **Prepare a Bell State**: Construct a quantum circuit that generates an EPR pair (Bell state $(\left|\Phi^+\right\rangle)$).
3. **Simulate the Circuit**: Use a simulator to execute the circuit with the specified `shots` and `seed`.
4. **Return Results**: Output the measurement results (counts).

### Parameters
- **`input_state`** (`list`): The initial quantum state vector of the single qubit (e.g., `[1, 0]` for |0⟩ or `[0, 1]` for |1⟩).
- **`shots`** (`int`): The number of measurement shots for the simulation.
- **`seed`** (`int`): The seed for the random number generator to ensure reproducibility.

### Expected Functionality
1. Build a quantum circuit that initializes a single qubit to `input_state`.
2. Use the initial state as input to a circuit that implements the construction of the specified Bell State.
3. Simulate the circuit using a quantum simulator.
4. Return the measurement results (counts).


```python
!pip install qiskit_aer
!pip install qiskit
```

```python
import numpy as np
from qiskit_aer import AerSimulator
from qiskit import QuantumCircuit, transpile
from qiskit.quantum_info import Statevector
from qiskit.visualization import plot_histogram, plot_state_city

def normalize(v):
    norm = np.linalg.norm(v)
    if norm == 0: 
       return v
    return v / norm

def create_EPRPair(input_state, shots, seed):

    qc = QuantumCircuit(2)
    qc.initialize(Statevector(input_state), [0])
    qc.h(0)
    qc.cx(0, 1)
    qc.measure_all()
   
    backend = AerSimulator()
    
    transpiled_qc = transpile(qc, backend)

    job = backend.run(transpiled_qc, shots=shots, seed_simulator=seed)

    result = job.result()
    counts = result.get_counts(transpiled_qc)
    print(counts)


    return counts

```

```python

```

Output:
```
{'00': 480, '11': 520}
{'00': 480, '11': 520}
```

<!-- grade_id: Problem2_Test1, points: 10, grade: true, solution: false -->
```python
''' test 1 '''
input_state = [0.6, 0.8j]
expected_counts = {'00': 480, '11': 520}  # Expected output from the function
counts = create_EPRPair(input_state, shots=1000, seed=10)
assert counts == expected_counts, f"Test failed"

```

Output:
```
{'00': 480, '11': 520}
```

<!-- grade_id: Problem2_Test3, points: 30, grade: true, solution: false -->
```python


```

Output:
```
┌──────────────────────┐
q_0: ┤ Initialize(0.6,0.8j) ├
     └──────────────────────┘
q_1: ────────────────────────
```