#### PROBLEMA 19
    Write a function quantum_or_gate(input1, input2) that simulates a classical OR gate using a quantum circuit. The circuit should take the input state ∣input1,input2,0⟩ and produce the output state ∣input1,input2,(input1 OR input2)⟩. The function should:
    - Initialize the qubits based on the binary values of input1 and input2.
    - Apply the appropriate quantum gate/s to compute the OR operation.
    - Combine all operations into a single gate and return the resulting quantum circuit.
    Note: Do not include measurement operations. 


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-27808da576c0e8c3, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit import QuantumCircuit
from qiskit.circuit.library import UnitaryGate
from qiskit.quantum_info import Operator

def quantum_or_gate(input1, input2):
   
    qc = QuantumCircuit(3)


    if input1 == 1:
        qc.x(0)  
    if input2 == 1:
        qc.x(1)  
    qc.ccx(0, 1, 2) 

    return qc
```

<!-- grade_id: cell-dfb4585baec78103, points: 25, grade: true, solution: false -->
```python
exp1 = [[(1+0j),0j,0j,0j,0j,0j,0j,0j,],[0j,0j,0j,0j,0j,(1+0j),0j,0j,],[0j,0j,0j,0j,0j,0j,(1+0j),0j,],[0j,0j,0j,0j,0j,0j,0j,(1+0j),],[0j,0j,0j,0j,(1+0j),0j,0j,0j,],[0j,(1+0j),0j,0j,0j,0j,0j,0j,],[0j,0j,(1+0j),0j,0j,0j,0j,0j,],[0j,0j,0j,(1+0j),0j,0j,0j,0j]]
out1 = quantum_or_gate(0,0)
assert np.allclose(exp1, out1), f"Test failed"
```

Output:
```
File ~/Enviroments/HEnv1/lib/python3.12/site-packages/numpy/_core/numeric.py:2443, in isclose(a, b, rtol, atol, equal_nan)
   2351 """
   2352 Returns a boolean array where two arrays are element-wise equal within a
   2353 tolerance.
   (...)
   2440 
   2441 """
   2442 # Turn all but python scalars into arrays.
-> 2443 x, y, atol, rtol = (
   2444     a if isinstance(a, (int, float, complex)) else asanyarray(a)
   2445     for a in (a, b, atol, rtol))
   2447 # Make sure y is an inexact type to avoid bad behavior on abs(MIN_INT).
   2448 # This will cause casting of x later. Also, make sure to allow subclasses
   2449 # (e.g., for numpy.ma).
   (...)
   2452 #       timedelta works if `atol` is an integer or also a timedelta.
   2453 #       Although, the default tolerances are unlikely to be useful
   2454 if (dtype := getattr(y, "dtype", None)) is not None and dtype.kind != "m":

File ~/Enviroments/HEnv1/lib/python3.12/site-packages/numpy/_core/numeric.py:2444, in <genexpr>(.0)
  
```

<!-- grade_id: cell-1dcfe111064f6665, points: 25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-aa32611f42608f85, points: 25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-94d47ca5e8b47aa1, points: 25, grade: true, solution: false -->
```python

```