#### PROBLEMA 18
    Create a function normalize_state(state_vector) that accepts an array of complex numbers representing the coefficients of an unnormalized quantum state for multiple qubits. The function should normalize the state vector such that the sum of the squared magnitudes of all coefficients equals 1, and return the normalized state vector.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-4025a8d1d540df25, points: 0, grade: false, solution: true -->
```python
import numpy as np

def normalize_state(state_vector):
    
    norm = np.linalg.norm(state_vector)
    normalized_state = state_vector / norm
    
    return normalized_state

```

<!-- grade_id: cell-de7da6c3418485c6, points: 20, grade: true, solution: false -->
```python
exp1 = [0.70710678, 0.70710678]
state_vector1 = np.array([1, 1])
out1 = normalize_state(state_vector1)
assert np.allclose(out1, exp1), f"Test failed"
```

<!-- grade_id: cell-e8e57d3ad41ba736, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-b026addb94af19dd, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-a42825c0d471c3be, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-144304000fdc9fb6, points: 20, grade: true, solution: false -->
```python

```