#### PROBLEMA 16
    Write a function bit_flip_code(input_state, error) that implements a repetition code for detecting bit-flip errors. The function should:
        - Encode a single-qubit quantum state (input_state) into a three-qubit (repetition code) state using a stabilizer code.
        - Apply bit-flip errors (X gates) based on the qubit indices provided in the error list.
        - Decode the three-qubit state and return the error syndrome after the decoding process.
    The function should use stabilizer code principles to detect bit-flip errors.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-01ca43c24ef9ff75, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit import QuantumCircuit, transpile
from qiskit.quantum_info import Statevector
from qiskit_aer import AerSimulator

def bit_flip_code(input_state, error_qubits):
    
    qc = QuantumCircuit(3, 2)
    qc.initialize(Statevector(input_state), [0])
    
    # encode
    qc.cx(0, 1)
    qc.cx(0, 2)
    
    for qubit in error_qubits:
        qc.x(qubit)

    # syndrome
    qc.cx(1, 2)
    qc.cx(0, 1)
    qc.measure([1, 2], [0, 1])
    
    # decode
    qc.cx(0, 2)
    qc.cx(0, 1)

    backend = AerSimulator()
    transpiled_qc = transpile(qc, backend)
    result = backend.run(transpiled_qc, shots=1).result()
    counts = result.get_counts(transpiled_qc)

    return next(iter(counts.keys()))

```

<!-- grade_id: cell-b0cb92ca5cfb544c, points: 40, grade: true, solution: false -->
```python
exp1 = '01'
input_state = [1/np.sqrt(2), 1/np.sqrt(2)]
out1 = bit_flip_code(input_state, [0])
assert exp1 == out1, f"Test failed"
```

<!-- grade_id: cell-fcf220a5d5bcdd59, points: 40, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-8ebb8c43b9ddad86, points: 40, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-bc28ffb39bde6230, points: 40, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-0317899dd0abf8d5, points: 40, grade: true, solution: false -->
```python

```