### PROBLEMA 32
    Write a function implement_mod2_oracle(num_qubits) that constructs a quantum oracle for the function f(x)=xmod  2. The function should take the number of input qubits (num_qubits) as a parameter and create a quantum circuit where the input state ∣x0⟩ is transformed into ∣x (xmod  2)⟩, with the result stored in an ancilla qubit. The function should return the quantum circuit implementing this oracle. Construct just the circuit that would implement this functionality.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-1a33b47a04a5d05d, points: 0, grade: false, solution: true -->
```python
from qiskit import QuantumCircuit
from qiskit.quantum_info import Operator
import numpy as np

def implement_mod2_oracle(num_qubits):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-70d550916739d403, points: 50, grade: true, solution: false -->
```python
exp1 = [[(1+0j),0j,0j,0j,0j,0j,0j,0j,],[0j,(1+0j),0j,0j,0j,0j,0j,0j,],[0j,0j,0j,0j,0j,0j,(1+0j),0j,],[0j,0j,0j,0j,0j,0j,0j,(1+0j),],[0j,0j,0j,0j,(1+0j),0j,0j,0j,],[0j,0j,0j,0j,0j,(1+0j),0j,0j,],[0j,0j,(1+0j),0j,0j,0j,0j,0j,],[0j,0j,0j,(1+0j),0j,0j,0j,0j]]
out1 = implement_mod2_oracle(2)
assert np.allclose(out1, exp1), f"Test failed"
```

<!-- grade_id: cell-d6ef329119097a82, points: 50, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-ae07cd2d80a36f2f, points: 50, grade: true, solution: false -->
```python

```