#### PROBLEMA 20
    Write a function measure_in_basis(state_vector, basis_matrix) that measures a qubit initially in the computational basis in an arbitrary basis defined by basis_matrix. The matrix represents the basis of transformation. The function should:
        1. Verify that basis_matrix is unitary; if not, raise an error with the message 'matrix is not unitary'.
        2. Compute and return the probabilities of measurement outcomes in the specified basis as an array.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-98319709080d8d73, points: 0, grade: false, solution: true -->
```python
import numpy as np

def measure_in_basis(state_vector, basis_matrix):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7818584f4ceea0e8, points: 20, grade: true, solution: false -->
```python
exp1 = [0.4999999999999999, 0.4999999999999999]
state = np.array([0.6, -0.8j])
X_obs_gate = (1/np.sqrt(2)) * np.array([[1, 1],[1, -1]])
out1 = measure_in_basis(state, X_obs_gate)
assert np.allclose(exp1, out1), f"Test failed"
```

<!-- grade_id: cell-78f9dfd3cce189f2, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-3ab5f59e7a7e14b3, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-66a0bd069e5b4e44, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-77073002ecd639e6, points: 20, grade: true, solution: false -->
```python

```