### PROBLEMA 24
    Write a Python function multi_controlled_not(num_qubits, control_qubits, target_qubit) that constructs and returns the matrix representation of a multi-controlled NOT gate. The function should:
        1. Accept num_qubits as the total number of qubits in the system.
        2. Take control_qubits as a list of indices representing the control qubits.
        3. Specify target_qubit as the index of the target qubit.
    Return the resulting matrix in the standard format of quantum computing (not Qiskit's representation).


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-d810bfe04f4b187a, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit import QuantumCircuit
from qiskit.quantum_info import Operator
# from qiskit.circuit.library import.


def multi_controlled_not(num_qubits, control_qubits, target_qubit):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-22503623e56f106c, points: 20, grade: true, solution: false -->
```python
exp1 = [[(1+0j),0j,0j,0j,],[0j,(1+0j),0j,0j,],[0j,0j,0j,(1+0j),],[0j,0j,(1+0j),0j]]
out1 = multi_controlled_not(2, [0], 1)
assert np.allclose(out1, exp1), f"Test failed"
```

<!-- grade_id: cell-ab5177486358a3ca, points: 20, grade: true, solution: false -->
```python

```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[1], line 2
      1 # Define the quantum circuit with 4 qubits
----> 2 qc = QuantumCircuit(4, 2)  # 4 qubits and 2 classical bits for measurement
      4 # Initialize inputs A=1, B=0, CarryIn=1
      5 qc.x(0)  # A = 1

NameError: name 'QuantumCircuit' is not defined
```

<!-- grade_id: cell-116c9eb621925770, points: 20, grade: true, solution: false -->
```python


```

<!-- grade_id: cell-c87fef63ea9998ee, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-1b71b8d401107d3c, points: 20, grade: true, solution: false -->
```python

```