Write a function inner_product(state1, state2) that calculates the inner product ⟨state1∣state2⟩ between two quantum states, represented as arrays of complex coefficients. The function should:
Verify that both states are normalized; raise an error with the message 'state is not normalized' if either state is not normalized.
Ensure that the states have the same dimension; raise an error with the message 'the vectors do not have the same dimension' if they differ.
Compute and return the inner product.
The input states are provided as row vectors.


```python
!pip install numpy
```

<!-- grade_id: cell-af0c142c3534d231, points: 0, grade: false, solution: true -->
```python
import numpy as np

def inner_product(state1, state2):
    # Check if both states are normalized
    if not np.isclose(np.linalg.norm(state1), 1):
        raise ValueError("state is not normalized")
    if not np.isclose(np.linalg.norm(state2), 1):
        raise ValueError("state is not normalized")
    
    # Check if both states have the same dimension
    if state1.shape != state2.shape:
        raise ValueError("the vectors do not have the same dimension")
    
   
    return np.vdot(state1, state2)
```

<!-- grade_id: cell-43a84ad591c960c4, points: 20, grade: true, solution: false -->
```python
out_exp = 0.0
state1_1 = (1/np.sqrt(2))*np.array([1, 1])
state1_2 = (1/np.sqrt(2))*np.array([1, -1])
out = inner_product(state1_1, state1_2)
print(out)
assert np.isclose(out, out_exp), f"Test failed"
```

Output:
```
0.0
```

<!-- grade_id: cell-43129e5b754f9206, points: 20, grade: true, solution: false -->
```python


```

<!-- grade_id: cell-4eeef74f76304669, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-c0561cbc65106beb, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-5f2dc1d6e2e35ee6, points: 20, grade: true, solution: false -->
```python

```