### PROBLEMA 22
    Write a function create_u3_gate(theta, phi, lambda) that generates the matrix representation of a custom U3 gate based on the input rotation angles θ, ϕ, and λ. The function should compute and return the U3 gate matrix. Use the normal way of representing qubits (not the one Qiskit does).


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-a2bfc8092600b8f3, points: 0, grade: false, solution: true -->
```python
import numpy as np


def create_u3_gate(theta, phi, lamda):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4c5d947fba523068, points: 20, grade: true, solution: false -->
```python
exp1 = [[ 1.+0.j, -0.+0.j], 
        [ 0.+0.j,  1.+0.j]]
theta1, phi1, lamda1 = 0, 0, 0
out1 = create_u3_gate(theta1, phi1, lamda1)
assert np.allclose(exp1, out1), f"Test failed"
```

<!-- grade_id: cell-77645620011b9a2a, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-f4f3af435df5ad15, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-ff87b56a67857b14, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-3cc09fa0521388fc, points: 20, grade: true, solution: false -->
```python

```