Write a function prepare_ghz_state(num_qubits, shots, seed) that generates a generalized ∣GHZ⟩ state for a given number of qubits (num_qubits). The function should take the number of measurements (shots) and a seed (seed) for reproducibility. After preparing the GHZ state, simulate the circuit, perform the specified number of measurements, and return the measurement results (counts).


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-9c744b72699b9c24, points: 0, grade: false, solution: true -->
```python
import numpy as np
from qiskit_aer import AerSimulator
from qiskit import QuantumCircuit, transpile


def prepare_ghz_state(num_qubits, shots, seed):
    qc = QuantumCircuit(3)
    qc.h(0)
    qc.cx(0, 1)
    qc.cx(0, 2)
    qc.measure_all()

    backend = AerSimulator()
    
    transpiled_qc = transpile(qc, backend)

    job = backend.run(transpiled_qc, shots=shots, seed_simulator=seed)

    result = job.result()
    counts = result.get_counts(transpiled_qc)
    print(counts)


    return counts
 
```

<!-- grade_id: cell-f76f3e4f8af93244, points: 24, grade: true, solution: false -->
```python
counts_expected = {'000': 5050, '111': 4950}
cts = prepare_ghz_state(3, 10000, 90)
assert cts == counts_expected, f"Test Failed"
```

Output:
```
{'111': 4950, '000': 5050}
```

<!-- grade_id: cell-f961343545eb3f66, points: 24, grade: true, solution: false -->
```python

```

Output:
```
┌───┐           ░ ┌─┐      
   q_0: ┤ H ├──■────■───░─┤M├──────
        └───┘┌─┴─┐  │   ░ └╥┘┌─┐   
   q_1: ─────┤ X ├──┼───░──╫─┤M├───
             └───┘┌─┴─┐ ░  ║ └╥┘┌─┐
   q_2: ──────────┤ X ├─░──╫──╫─┤M├
                  └───┘ ░  ║  ║ └╥┘
meas: 3/═══════════════════╩══╩══╩═
                           0  1  2
```

<!-- grade_id: cell-c4f456e12733016a, points: 24, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-a4fa4da19a67efb9, points: 24, grade: true, solution: false -->
```python

```

```python

```