# Quantum Teleportation
**Challenge:** 

> Implement a function ```quantum_teleportation(input_state, seed)``` that constructs a quantum circuit to perform quantum teleportation and returns the result counts.

*Details:*

1. Initialize the Circuit: Create a quantum circuit with three qubits.
2. Prepare Bell State: Apply Hadamard and CNOT gates to the first two qubits to create an entangled Bell state.
3. Entangle and Measure: Entangle the input state with the first qubit and measure the first two qubits.
4. Apply Corrections: Apply appropriate X and Z gates to the third qubit based on the measurement results to complete the teleportation.


```python
!pip install qiskit[visualization]
!pip install qiskit-aer
!pip install numpy
```

```python
from qiskit import QuantumCircuit, transpile
from qiskit_aer import Aer
import numpy as np
from classiq import *
```

<!-- grade_id: cell-78642679ee29cf44, points: 0, grade: false, solution: true -->
```python
def quantum_teleportation(input_state, seed):
    # Initialize the circuit
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Prepare Bell state
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Entangle and measure
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Apply conditional gates based on the measurement results
    # Apply Z gate if the first classical bit is 1
    # Apply X gate if the second classical bit is 1
    # YOUR CODE HERE
    raise NotImplementedError()

    # Measure the teleported qubit
    circuit.measure(2, 2)

    # Transpile and assemble the circuit for the backend
    backend = Aer.get_backend('qasm_simulator')
    transpiled_circuit = transpile(circuit, backend)

    # Run the transpiled circuit on the backend and return the counts
    job = backend.run(transpiled_circuit, shots=1024, seed_simulator=seed)
    result = job.result()
    counts = result.get_counts()
    return counts
```

<!-- grade_id: cell-4df107be8bda1aa4, points: 100, grade: true, solution: false -->
```python
"""Check that the function returns the correct output for several inputs"""
assert quantum_teleportation([1/np.sqrt(2), -1/np.sqrt(2)], 1234) == {'110': 135, '100': 118, '101': 112, '000': 124, '011': 128, '001': 138, '010': 132, '111': 137}
```