Write a function is_quantum_gate(matrix) that determines whether a given matrix represents a valid quantum gate. The function should return True if the matrix satisfies the necessary conditions for a quantum gate and False otherwise.

```python
!pip install numpy
```

<!-- grade_id: Problema8, points: 0, grade: false, solution: true -->
```python
import numpy as np
def is_quantum_gate(matrix):
    conjugate_transpose = matrix.conj().T
    inverse = np.linalg.inv(matrix)
    return np.array_equal(conjugate_transpose, inverse)
    

```

<!-- grade_id: Problema8_T1, points: 4, grade: true, solution: false -->
```python
''' test 1 '''
exp = True
matrix = np.array([[1, 0], [0, 1]])
out = is_quantum_gate(matrix)
assert exp == out, f"Test failed"

```

<!-- grade_id: Problema8_T3, points: 24, grade: true, solution: false -->
```python

```

<!-- grade_id: Problema8_T4, points: 24, grade: true, solution: false -->
```python

```

<!-- grade_id: Problema8_T5, points: 24, grade: true, solution: false -->
```python

```

```python

```