# Grover's Algorithm
**Challenge:**

> Implement a function ```grovers_algorithm(target_state, seed)``` that constructs a quantum circuit to perform Grover's search algorithm and returns the result.

**Details:**

1. Initialize the Circuit: Create a quantum circuit with enough qubits to represent the search space and an additional qubit for the ancilla.
2. Apply Hadamard Gates: Place all qubits into a superposition state by applying Hadamard gates.
3. Oracle Implementation: Create an oracle that flips the amplitude of the target state. This can be done using a multi-controlled Toffoli gate.
4. Diffusion Operator: Apply the diffusion operator to amplify the probability of the target state.
5. Iteration: Repeat the application of the oracle and diffusion operator √N times (where N is the number of states).
6. Measurement: Measure the qubits to find the target state.

```python
!pip install qiskit[visualization]
!pip install qiskit-aer
!pip install numpy
!pip install classiq
```

```python
from qiskit import QuantumCircuit, transpile
from qiskit_aer import Aer
from qiskit.visualization import plot_histogram
import numpy as np
from classiq import *
```

<!-- grade_id: cell-d0e9f36ca1f53481, points: 0, grade: false, solution: true -->
```python
def grovers_algorithm(target_state, seed):
    # Initialize the circuit
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Oracle implementation
    def apply_oracle(circuit, target_state):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    # Diffusion operator
    def apply_diffusion_operator(circuit):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    # Grover iterations
    for _ in range(int(np.sqrt(2 ** num_qubits))):
        circuit = apply_oracle(circuit, target_state)
        circuit = apply_diffusion_operator(circuit)
        
    
    # Measurement
    circuit.measure([0, 1, 2], [0, 1, 2])
    backend = Aer.get_backend('qasm_simulator')
    transpiled_circuit = transpile(circuit, backend)
    job = backend.run(transpiled_circuit, shots=1024, seed_simulator=seed)
    result = job.result()
    counts = result.get_counts()
    return counts
    
```

<!-- grade_id: cell-3f91ffc59155f87a, points: 80, grade: true, solution: false -->
```python
"""Check that the function returns the correct output for several inputs"""
assert grovers_algorithm('110', 42) == {'110': 983, '011': 5, '111': 7, '100': 7, '101': 9, '000': 6, '001': 3, '010': 4}
```

<!-- grade_id: cell-9a916f7579924e5f, points: 20, grade: true, solution: false -->
```python
"""Check for invalid inputs"""
```