### PROBLEMA 30
    Create a function bloch_sphere_angles(state_vector) that computes the angles required for representing a single-qubit quantum state on the Bloch sphere. The function should:
1. Validate that the input state is normalized, and raise an error with the message 'state is not normalized' if it is not.
2. Return the polar angle θ, the azimuthal angle ϕ, and the global phase ω as an array.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-6784388b39fabc9e, points: 0, grade: false, solution: true -->
```python
import numpy as np

def bloch_sphere_angles(state_vector):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-af78c147a4e7fc35, points: 20, grade: true, solution: false -->
```python
exp1 = (1.5707963267948968, 0.0, 0.0)
state_vector1 = np.array([1/np.sqrt(2), 1/np.sqrt(2)])
out1 = bloch_sphere_angles(state_vector1)
assert np.allclose(out1, exp1), f"Test failed"
```

<!-- grade_id: cell-5f7f58d13a7c31d7, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-f3e6404aa060bbdc, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-db4df458553eb4bc, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-0d50ddd55d7575c2, points: 20, grade: true, solution: false -->
```python

```