# Quantum conditional statements with Classiq

**Challenge:** Implement a function ```q_if(initial_state, test_value)``` that creates a circuit to check if the initial state is equal to the ```test_value``` without collapsing the state. The function should return the results of the measurement which will highlight with a target qubit the value 1 if the initial state matches the test_value, or 0 otherwise.


```python
!pip install classiq
```

```python
from classiq import *
```

<!-- grade_id: cell-fda87242456ad3bd, points: 0, grade: false, solution: true -->
```python
def q_if(initial_state, test_value):
    
    @qfunc
    def main(x: Output[QNum], target: Output[QBit]):
        # YOUR CODE HERE
        raise NotImplementedError()

    # Create the quantum program and execute the job
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Get the results
    results = job.result()[0].value.parsed_counts

    # Get the data of the circuit
    data = QuantumProgram.from_qprog(quantum_program).data

    # Return the results
    return results, data
```

<!-- grade_id: cell-92f6e0fafceb3b69, points: 0, grade: true, solution: false -->
```python
"""Check that the function returns the correct output for several inputs"""

assert "'target': 1" in f'{q_if(3, 3)[0]}'
assert "'target': 0" in f'{q_if(3, 2)[0]}'
```