Write a function normalize_state(alpha, beta) that accepts two complex numbers, alpha and beta, representing the coefficients of an unnormalized quantum state ∣ψ⟩=α∣0⟩+β∣1⟩. The function should return the normalized state ∣ψ′⟩=α′∣0⟩+β′∣1⟩ (ensuring that ∣α′∣2+∣β′∣2=1) as an array.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-fd0b1893e14812fc, points: 0, grade: false, solution: true -->
```python
import numpy as np

def normalize_state(alpha, beta):
    norm = np.sqrt(abs(alpha)**2 + abs(beta)**2)
    
    # Normalize the coefficients
    alpha_normalized = alpha / norm
    beta_normalized = beta / norm
    return [alpha_normalized, beta_normalized]
  
```

<!-- grade_id: cell-d66b9ef75486a880, points: 10, grade: true, solution: false -->
```python
out_expected = [0.6, 0.8]
out = normalize_state(3, 4)
print(out)
assert np.allclose(out, out_expected), f"Test Failed"

```

Output:
```
[np.float64(0.6), np.float64(0.8)]
```

<!-- grade_id: cell-84f70c996a70b161, points: 22.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-96df8c2572755c67, points: 22.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-59b2adfe9397ecbc, points: 22.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-7826e78c1b9174fa, points: 22.5, grade: true, solution: false -->
```python

```