#### PROBLEMA 21
    Write a function transform_to_basis(input_state, new_basis) that transforms a quantum state vector from the computational basis to a specified new basis, which is given by a list of its eigenvectors The function should:
        - Verify that the vectors in the list new_basis are orthonormal; if not, raise an error with the message 'new basis is not orthonormal'.
        - Transform the input_state into the new_basis using the provided eigenvectors.
        - Return the transformed state vector in the new basis as an array
    Note: Be aware that the eigenvectors are obtained starting from some observables, using the specific function in numpy. Thus, the order of the eigenvectors could be reversed than the normal way of writing transformation matrices (we will use common ones)


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-23ecb95c4768ff1d, points: 0, grade: false, solution: true -->
```python
import numpy as np

def transform_to_basis(input_state, new_basis):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5a4717b494b4ccff, points: 10, grade: true, solution: false -->
```python
exp1 = [1., 0.]
state1 = np.array([1/np.sqrt(2), 1/np.sqrt(2)])
basis_matrix1 = np.linalg.eig(np.array([[0, 1], [1, 0]]))[1]
out1 = transform_to_basis(state1, basis_matrix1)
assert np.allclose(exp1, out1), f"Test failed"
```

<!-- grade_id: cell-e800a91f452a6864, points: 18, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-e2f5b25709e4b452, points: 18, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-e4a9e1e98bb48561, points: 18, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-0e233890adeaf171, points: 18, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-2b6d57cff01d7f21, points: 18, grade: true, solution: false -->
```python

```