### PROBLEMA 31
    Create a function toffoli_with_basic_gates(num_qubits, control1, control2, target) that constructs a Toffoli (CCX) gate using only fundamental quantum gates such as CNOT and single-qubit gates. The function should represent the construction in the standard, paper-based notation of quantum gates (not specific to Qiskit) and encapsulate the resulting circuit into a single gate. Return the final constructed gate.


```python
!pip install numpy
!pip install qiskit
!pip install qiskit_aer
```

<!-- grade_id: cell-7e08ded47b5c9a91, points: 0, grade: false, solution: true -->
```python
from qiskit import QuantumCircuit
from qiskit.quantum_info import Operator
import numpy as np

def toffoli_with_basic_gates(num_qubits, control1, control2, target):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8ff84a7c18bc8026, points: 40, grade: true, solution: false -->
```python
exp1 = [[(0.9999999999999997+0j),(1.1102230246251565e-16+0j),0j,0j,0j,0j,0j,0j,],[(1.1102230246251565e-16+0j),(0.9999999999999997+0j),0j,0j,0j,0j,0j,0j,],[0j,0j,(0.9999999999999996+0j),0j,0j,0j,0j,0j,],[0j,0j,0j,(0.9999999999999996+0j),0j,0j,0j,0j,],[0j,0j,0j,0j,(0.9999999999999996+0j),0j,0j,0j,],[0j,0j,0j,0j,0j,(0.9999999999999996+0j),0j,0j,],[0j,0j,0j,0j,0j,0j,0j,(0.9999999999999996+0j),],[0j,0j,0j,0j,0j,0j,(0.9999999999999996+0j),0j]]
out1 = toffoli_with_basic_gates(3, 0, 1, 2)
assert np.allclose(out1, exp1), f"Test failed"
```

<!-- grade_id: cell-097db8609765d8dd, points: 40, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-d81f3243193ad958, points: 40, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-64cf7e55a8b42c2a, points: 40, grade: true, solution: false -->
```python

```